rule_files:
  - ../rules/alerts.yml

evaluation_interval: 1m

tests:
- interval: 1m
  input_series:
    - series: 'http_requests_total:counter{endpoint="/api/auth/login",status="401"}'
      values: '0 10 20 30 40 50 60 70 80 90 100'
  alert_rule_test:
    - eval_time: 10m
      alertname: Login401429Spike
      exp_alerts:
        - exp_labels:
            alertname: Login401429Spike
            severity: warning
- interval: 1m
  input_series:
    - series: 'av_scan_failures_total:counter{provider="clamav"}'
      values: '0 0 0 0 1 1 1'
  alert_rule_test:
    - eval_time: 6m
      alertname: AVScanFailure
      exp_alerts:
        - exp_labels:
            alertname: AVScanFailure
            severity: critical
            provider: clamav
          exp_annotations:
            summary: "Antivirus scan failure detected"
            description: "An antivirus scan has failed in the last 5 minutes"
- interval: 1m
  input_series:
    - series: 'http_requests_total:counter'
      values: '0 100 200 300 400 500 600 700 800 900 1000'
    - series: 'http_requests_errors_total:counter'
      values: '0 10 20 30 40 50 60 70 80 90 100'
  alert_rule_test:
    - eval_time: 10m
      alertname: HighErrorRate
      exp_alerts:
        - exp_labels:
            alertname: HighErrorRate
            severity: critical
- interval: 1m
  input_series:
    - series: 'http_request_duration_seconds_bucket:counter{le="0.5"}'
      values: '0 50 100 150 200 250 300 350 400 450 500'
    - series: 'http_request_duration_seconds_bucket:counter{le="1"}'
      values: '0 90 180 270 360 450 540 630 720 810 900'
    - series: 'http_request_duration_seconds_bucket:counter{le="3"}'
      values: '0 100 200 300 400 500 600 700 800 900 1000'
    - series: 'http_request_duration_seconds_bucket:counter{le="+Inf"}'
      values: '0 100 200 300 400 500 600 700 800 900 1000'
  alert_rule_test:
    - eval_time: 10m
      alertname: HighResponseTime
      exp_alerts:
        - exp_labels:
            alertname: HighResponseTime
            severity: warning
